---
title       : Basics of Julia Programming
description : In this Section we will begin to explore Julia and learn how it handles data. We will focus on major data types that we use frequently in statistics and the basics of how we use them. As we move forward these data types will be key to analyzing statistics.  
attachments :


--- type:VideoExercise lang:r xp:25 skills:1    key:8164d3d037
## Getting Start in Julia

*** =video_link

//player.vimeo.com/video/



--- type:NormalExercise lang:r xp:25 skills:1  key:d4bdee1c7f
##Practicing Basic Julia

*** =instructions

This course site is only equipped to handle either R or Python, because of this you will need to run the code for all Julia exercises on your own computer and copy the answers into R as a comment. 

Copy the following Julia output as a comment:

```
julia> 5+5
10
```



*** =hint

Here is a hint: Remember to use # for comments

*** =sample_code

```{r}
# Copy the code in as a comment




```

*** =solution

```{r}
# Copy the code in as a comment
#julia> 5+5
#10
```

*** =sct
```{r}
test_output_contains("julia", incorrect_msg = "Make sure to print `a`")
success_msg("Great!")
```


--- type:VideoExercise lang:r xp:25 skills:1    key:b29f203e7b
## Simple Procedures in Julia

*** =video_link

//player.vimeo.com/video/




--- type:VideoExercise lang:r xp:25 skills:1    key:c4dd2f996d
## Data Objects in Julia

*** =video_link

//player.vimeo.com/video/


--- type:VideoExercise lang:r xp:25 skills:1    key:b62fd8174d
## Arrays in Julia

*** =video_link

//player.vimeo.com/video/


--- type:VideoExercise lang:r xp:25 skills:1    key:22254b93ba
## Working with Arrays in Julia

*** =video_link

//player.vimeo.com/video/


